rule_files:
    - rules.yml

evaluation_interval: 1m

tests:

    - name: RDSLowDiskSpaceCount
      interval: 1m
      input_series:
          - series: 'rds_free_storage_bytes{aws_account_id="111111111111",aws_region="eu-west-3",dbidentifier="db1"}'
            values: '3221225472x10'  # 3GB
          - series: 'rds_allocated_storage_bytes{aws_account_id="111111111111",aws_region="eu-west-3",dbidentifier="db1"}'
            values: '21474836480x15'  # 20GB
      alert_rule_test:
          - alertname: RDSLowDiskSpaceCount
            eval_time: 15m
            exp_alerts:
                - exp_labels:
                      aws_account_id: 111111111111
                      aws_region: eu-west-3
                      severity: warning
                  exp_annotations:
                      description: "One or more RDS instances has <20% free disk space"
                      summary: "Less than 20% free disk space on at least one instance"
                      runbook_url: "https://qonto.github.io/database-monitoring-framework/0.0.0/runbooks/rds/RDSLowDiskSpaceCount"
